# Программирование для всех: основы Python

## Задачи для тренировки

*Алла Тамбовцева*

*Не сдаются на проверку, можно порешать для практики.<br>Позже будет опубликован разбор и примеры кода с решениями.*

Задачи A1-A6 считаются простыми (базовый уровень), задачи B1-B8 средними (начальный уровень), задачи C1-C2 – сложными (продвинутый уровень). Задача C1 несильно сложная, там надо просто догадаться об удобном преобразовании строки,  задача C2 – замороченная в плане структуры кода.

## Активируем тесты

```python
# Фабрика тестов для проверки программ, принимающих данные через input()

from collections import deque

class Tester(object):
    def __init__(self, inp):
        self.outputs = []
        self.inputs = deque(inp)
    def print(self, *args, sep = " ", end = "\n"):
        text = sep.join(map(str, args)) + end
        newlines = text.splitlines(keepends=True)
        if self.outputs and self.outputs[-1] and self.outputs[-1][-1] != "\n" and newlines:
            self.outputs[-1] += newlines[0]
            self.outputs.extend(newlines[1:])
        else:
            self.outputs.extend(newlines)
            
    def input(self, *args):
        assert self.inputs, "Вы пытаетесь считать больше элементов, чем предусмотрено условием"
        return self.inputs.popleft()
    def __enter__(self):
        global print
        global input
        print = self.print
        input = self.input
        return self.outputs
    def __exit__(self, *args):
        global print
        global input
        del print
        del input
```

### Задача А1. Power of numbers

Напишите программу, которая запрашивает с клавиатуры два целых положительных числа `a` и `b` (отдельный запрос для каждого значения) и выводит на экран строку следующего вида, с подставленными значениями на месте квадратных скобок:

    [a] в степени [b] равно [значение 1], [b] в степени [a] равно [значение 2]

**Примеры работы программы**

Ввод:

    2
    3

Вывод:

    2 в степени 3 равно 8, 3 в степени 2 равно 9

Ввод:

    4
    3

Вывод:

    4 в степени 3 равно 64, 3 в степени 4 равно 81

<!-- grade_id: ans04, points: 0, grade: false, solution: true -->
```python
def power_numbers():
    
    ### YOUR CODE HERE ###

power_numbers()
```

<!-- grade_id: test04, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("4 3", "4 в степени 3 равно 64, 3 в степени 4 равно 81"),
    ("2 5", "2 в степени 5 равно 32, 5 в степени 2 равно 25"),
    ("3 6", "3 в степени 6 равно 729, 6 в степени 3 равно 216"), 
    ("1 5", "1 в степени 5 равно 1, 5 в степени 1 равно 5"),
    ("1 1", "1 в степени 1 равно 1, 1 в степени 1 равно 1")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        power_numbers()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача А2. А не посчитать ли нам, уважаемые кроты...

Крот хочет вести учет количества рисовых зерен, которые съедает Дюймовочка. Напишите программу, которая:

* запрашивает с клавиатуры строку с дробными числами через пробел (одно значение – часть рисового зернышка, которое съедает Дюймовочка за день);

* выводит на экран суммарное число зерен с точностью до первого знака после точки.

**Примеры работы программы**

Ввод:

    0.5 0.6 0.4 0.2 0.3 0.5

Вывод:

    2.5

Ввод:

    0.2 0.5 0.5 0.3

Вывод:

    1.5

<!-- grade_id: cell-6ca594a009388498, points: 0, grade: false, solution: true -->
```python
def greedy_mole():

    ### YOUR CODE HERE ###
            
greedy_mole()
```

<!-- grade_id: cell-f1e56f11874a7ed2, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("0.5 0.6 0.4 0.2 0.3 0.5", "2.5"),
    ("0.2 0.5 0.5 0.3", "1.5"),
    ("0.3 0.3 0.4 0.8 0.6 0.5", "2.9")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        greedy_mole()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача А3. Поле большое, зелёный лесок...

Паровозик из Ромашкова – романтическая натура. Он часто опаздывает, потому что не может не обращать внимания на красоту вокруг: то рассветом залюбуется, то букет соберёт, то птенца выпавшего из гнезда поднимет. Правда, приоритеты у него и пассажиров разные.

Реализуйте программу, которая угадывает, кто вы:

* Дан список `parovozik` с приоритетами паровозика.
* Дан список `passenger` с приоритетами пассажиров.
* В переменную `prior` считывается приоритет пользователя. Этот код уже написан.
* Если введённый приоритет совпадает с одним из приоритетов паровозика, программа выводит строку `"Вы - паровозик!"`. Если приоритет совпадает с одним из приоритетов пассажиров, программа выводит строку `"Вы - пассажир!"`. Если ни то, ни другое, программа выводит строку `"Диспетчер?"`.

**Примеры работы программы**

Ввод:

    собрать ландыши


Вывод:

    Вы - паровозик!
    
Ввод:

    ехать быстро


Вывод:

    Вы - пассажир!
    
Ввод:

    не уснуть


Вывод:

    Диспетчер?

```python
def check_priorities():

    parovozik = ["собрать ландыши", "послушать пение соловья", "увидеть рассвет"]
    passenger = ["успеть вовремя", "ехать быстро"]
    prior = input()
    
    ### YOUR CODE HERE ###

check_priorities()
```

```python
test_data = [
    ("собрать ландыши", "Вы - паровозик!"),
    ("послушать пение соловья", "Вы - паровозик!"),
    ("увидеть рассвет", "Вы - паровозик!"),
    ("успеть вовремя", "Вы - пассажир!"),
    ("ехать быстро", "Вы - пассажир!"),
    ("не уснуть", "Диспетчер?"),
    ("something", "Диспетчер?"),
    ("", "Диспетчер?")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        check_priorities()
        assert len(t) == 1, "Неверный ответ"
        assert t[0] == out+"\n", "Неверный ответ"
```

### Задача А4. Жду ответа

У питона Пети в записной книжке хранятся имена друзей из разных стран. Записная книжка представляет собой список кортежей `pythons`.

    pythons = [("королевский питон Рональд", "Ангола"), ("карликовый питон Антониа", "Ангола"), 
               ("сетчатый питон Асиф", "Индонезия"), ("зеленый питон Оливер", "Австралия"), 
               ("тиморский питон Николау", "Китай"), ("аметистовый питон Харрисон", "Австралия"), 
               ("тигровый питон Мэйли", "Китай"), ("оливковый питон Руби", "Австралия"),
               ("белогубый питон Арис", "Индонезия"), ("эфиопский питон Оника", "Мозамбик")]

Пете очень захотелось съездить в какую-нибудь экзотическую страну. А для этого нужно заранее написать или позвонить друзьям или знакомым, чтобы выяснить, можно ли у них остановиться.

Напишите программу, которая запрашивает с клавиатуры название страны и выводит на экран список питонов из этой страны в том порядке, в котором они встречаются в записной книжке.

Уточнения:

* Названия стран вводятся с клавиатуры в том виде, в котором они сохранены в записной книжке `pythons`.
* Если введено название страны, которого в записной книжке нет, на экран выводится пустой список.

**Примеры работы программы**

Ввод:

    Австралия

Вывод:

    ['зеленый питон Оливер', 'аметистовый питон Харрисон', 'оливковый питон Руби']

Ввод:

    Китай

Вывод:

    ['тиморский питон Николау', 'тигровый питон Мэйли']
    
Ввод:

    Франция

Вывод:

    []

```python
def wait_answer():
    
    pythons = [("королевский питон Рональд", "Ангола"), ("карликовый питон Антониа", "Ангола"), 
               ("сетчатый питон Асиф", "Индонезия"), ("зеленый питон Оливер", "Австралия"), 
               ("тиморский питон Николау", "Китай"), ("аметистовый питон Харрисон", "Австралия"), 
               ("тигровый питон Мэйли", "Китай"), ("оливковый питон Руби", "Австралия"),
               ("белогубый питон Арис", "Индонезия"), ("эфиопский питон Оника", "Мозамбик")]

    ### YOUR CODE HERE ###
    
wait_answer()
```

```python
test_data = [
    ("Австралия", "['зеленый питон Оливер', 'аметистовый питон Харрисон', 'оливковый питон Руби']"),
    ("Индонезия", "['сетчатый питон Асиф', 'белогубый питон Арис']"),
    ("Ангола", "['королевский питон Рональд', 'карликовый питон Антониа']"), 
    ("Китай", "['тиморский питон Николау', 'тигровый питон Мэйли']"),
    ("Франция", "[]"),
    ("Россия", "[]")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        wait_answer()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача А5. Just `for` fun

При проверке, являются ли числовые данные сфальсифицированными, часто обращают внимание на количество чисел, которые делятся на 5 или 10 без остатка, так как считается, что люди чисто психологически склонны выдумывать именно такие числа. 

Напишите, программу, которая:

* запрашивает с клавиатуры два целых числа через пробел (назовем их `n1` и `n2`);
* выводит на экран числа от `n1` до `n2` включительно, при этом на каждое из них через пробел выводит комментарий `yes` (если число делится на 5 или 10 без остатка) или `no` (если иначе);
* в конце выводит на экран кортеж с числом ответов `yes` и `no` вида (`"yes"`, `k1`, `"no"`, `k2`), где `k1` и `k2` – количество ответов `yes` и `no` соответственно.

**Примеры работы программы**

Ввод:

    10 18

Вывод:

    10 yes
    11 no
    12 no
    13 no
    14 no
    15 yes
    16 no
    17 no
    18 no
    ('yes', 2, 'no', 7)
    
    
Ввод:

    -5 6

Вывод:

    -5 yes
    -4 no
    -3 no
    -2 no
    -1 no
    0 yes
    1 no
    2 no
    3 no
    4 no
    5 yes
    6 no
    ('yes', 3, 'no', 9)

```python
def just_for():
    
    ### YOUR CODE HERE ###
    
just_for()
```

```python
test_data = [("-5 6", ["-5 yes\n", "-4 no\n", "-3 no\n", "-2 no\n", 
                      "-1 no\n", "0 yes\n", "1 no\n", "2 no\n", "3 no\n",
                      "4 no\n", "5 yes\n", "6 no\n", "('yes', 3, 'no', 9)"]),
             ("1 9", ["1 no\n", "2 no\n", "3 no\n", "4 no\n", "5 yes\n", 
                      "6 no\n", "7 no\n", "8 no\n", "9 no\n", 
                      "('yes', 1, 'no', 8)\n"])]
for inp, out in test_data:
    with Tester(inp.split(",")) as t:
        just_for()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверный ответ"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), "Неверный ответ"
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Неверный ответ" 
```

### Задача A6. We are the champions...

Реализуйте программу, которая возвращает словарь с именем победителя по итогам серии игр и числом баллов, набранным победителем. Напишите функцию `champion()`, которая принимает на вход три аргумента:

* `results`: словарь с двумя записями, первая запись состоит из ключа `'best_name'` и имени победителя, вторая запись - из ключа `'best_score'` и числа набранных баллов;

* `name`: строка с именем участника;

* `score`: целое число баллов, набранное этим участником.

Если число очков, указанное в аргументе `score`, больше, чем число очков, сохраненное с ключом `'best_score'` в словаре, функция заменяет число очков в `'best_score'` на `score`, а имя в `'best_name'` – на `name`. Если иначе – функция возвращает исходный словарь `results` как есть.

ВАЖНО! Функция не должна ничего выводить на экран.

**Примеры работы функции**

Применение функции:
 
    champion({'best_name' : 'Alex', 'best_score': 75}, 'Alice', 42)


Результат:

    {'best_name' : 'Alex', 'best_score': 75}

Применение функции:
 
    champion({'best_name' : 'Alex', 'best_score': 75}, 'Anton', 92)


Результат:

    {'best_name' : 'Anton', 'best_score': 92}

```python
### YOUR CODE HERE ###
```

```python
assert champion({'best_name' : 'Alex', 'best_score': 75}, 
                'Anton', 42) == {'best_name': 'Alex', 'best_score': 75}, "Неверно"
assert champion({'best_name' : 'Alex', 'best_score': 75}, 
         'Alice', 92) == {'best_name': 'Alice', 'best_score': 92}, "Неверно"
assert champion({'best_name' : 'Alice', 'best_score': 63}, 
         'Sue', 40) == {'best_name': 'Alice', 'best_score': 63}, "Неверно"
assert champion({'best_name' : 'Sue', 'best_score': 44}, 
         'Victor', 98) == {'best_name': 'Victor', 'best_score': 98}, "Неверно"
```

### Задача B1. Я не жадный, я домовитый

Домовёнок Кузя ведёт учёт продуктов в холодильнике у Наташи. Напишите программу, которая поможет ему в этом.

* Программа запрашивает с клавиатуры названия продуктов через запятую с пробелом, одно слово соответствует одному продукту.
* Программа выводит на экран список кортежей с парами *продукт-количество*, упорядоченный по возрастанию количества продуктов.

**Примеры работы программы:**

Ввод:

    молоко, сосиска, сосиска, сосиска, огурец, огурец


Вывод:

    [('молоко', 1), ('огурец', 2), ('сосиска', 3)]
    
Ввод:

    яблоко, яблоко, масло, сырок, сырок, сырок, сырок, яблоко, сырок
    
Вывод:

    [('масло', 1), ('яблоко', 3), ('сырок', 5)]

```python
def fridge_inspector():
    
    ### YOUR CODE HERE ###
    
fridge_inspector()
```

```python
test_data = [
    ("молоко, сосиска, сосиска, сосиска, огурец, огурец", 
     "[('молоко', 1), ('огурец', 2), ('сосиска', 3)]"),
    ("яблоко, яблоко, масло, сырок, сырок, сырок, сырок, яблоко, сырок",
    "[('масло', 1), ('яблоко', 3), ('сырок', 5)]"),
    ("яйцо, яйцо, яйцо, молоко, йогурт, йогурт, котлета, котлета, яйцо",
    "[('молоко', 1), ('котлета', 2), ('йогурт', 2), ('яйцо', 4)]")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        fridge_inspector()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача B2. Они спешат на помощь

Дан словарь `heroes`. Добавьте в словарь `heroes` запись `Gadget` с информацией про Гайку: 

* `name`: Gadget Hackwrench;
* `attribute`: purple goggles;
* `inspiration`: NA.

Напишите программу, которая запрашивает с клавиатуры название характеристики и выводит ее значение для всех героев через запятую с пробелом с сортировкой по алфавиту.

**Пример работы программы**

Ввод:

    name

Вывод:

    Chip, Dale, Gadget Hackwrench
    
Ввод:

    attribute

Вывод:

    Hawaiian shirt, fedora hat, purple goggles
    
Ввод:

    inspiration

Вывод:

    Indiana Jones, NA, Thomas Magnum

```python
def rangers():
    
    heroes = {"Chip" : {"name" : 'Chip', 
                       "attribute" : 'fedora hat', 
                       "inspiration" : 'Indiana Jones'}, 
              "Dale" : {"name" : 'Dale', 
                       "attribute" : 'Hawaiian shirt',
                       "inspiration" : 'Thomas Magnum'}
             }
    
    ### YOUR CODE HERE ###
    
rangers()
```

```python
test_data = [
    ("name", "Chip, Dale, Gadget Hackwrench"),
    ("attribute", "Hawaiian shirt, fedora hat, purple goggles"), 
    ("inspiration", "Indiana Jones, NA, Thomas Magnum")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        rangers()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача B3. Ищи на третьей планете!

Как известно, птица Говорун отличается умом и сообразительностью. К тому же она знает, что искать пропавшего капитана Кима нужно на третьей планете. Реализуйте программу, сохраняющую в список номера планет, которые нужно посетить при поисках. 

Детали:

* В цикле считываются целые неотрицательные числа до тех пор, пока не введено 3 – это код нужной планеты. Гарантируется, что до 3 будет введено как минимум одно целое неотрицательное число.
* Для каждого вводимого числа проводится проверка: если число не равно 0 (это код Земли), оно добавляется в список планет, которые необходимо посетить. Если число уже есть в списке, повторно его добавлять не нужно.
* Список с номерами планет выводится на экран.

**Примеры работы программы:**

Ввод:

    2
    1
    3

Вывод:

    [2, 1]
    
Ввод:

    2
    1
    2
    2
    1
    3

Вывод:

    [2, 1]
    
Ввод:

    0
    3

Вывод:

    []
    
    
Ввод:

    1
    2
    0
    3

Вывод:

    [1, 2]

```python
def search_planet():

    ### YOUR CODE HERE ###
    
search_planet()
```

```python
test_data = [("2,1,3", ["[2, 1]"]),
             ("2,1,2,2,1,3", ["[2, 1]"]),
             ("0,3", ["[]"]),
             ("1,1,2,0,3", ["[1, 2]"])
             ]
for inp, out in test_data:
    with Tester(inp.split(",")) as t:
        search_planet()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверный ответ"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), "Неверный ответ"
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Неверный ответ" 
```

### Задача B4.  И снова Дюймовочка

В словаре `heroes` сохранены имена героев сказки «Дюймовочка» и списки того, что эти герои любят. Напишите программу, которая запрашивает с клавиатуры сначала одно имя, затем – второе (каждое имя с новой строки), а затем выводит на экран список объектов, которые нравятся обоим героям одновременно, отсортированный по алфавиту.

**Примеры работы программы:**

Ввод:

    Крот
    Мышь
    
Вывод:
    
    ['деньги', 'уют']
    
Ввод:

    Ласточка
    Крот
    
Вывод:
    
    []

<!-- grade_id: cell-07398649173aad3b, points: 0, grade: false, solution: true -->
```python
def get_info():
    
    heroes = {"Крот" : ["деньги", "темнота", "уют"], 
              "Мышь" : ["уют", "деньги", "беспрекословность"], 
              "Ласточка" : ["солнце", "свет", "весна"], 
              "Дюймовочка" : ["солнце", "свет", "цветы"],
              "Принц эльфов" : ["цветы", "Дюймовочка"]
             }
    
    ### YOUR CODE HERE ###
    
get_info()
```

<!-- grade_id: cell-d2cd8692dd1118ab, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("Ласточка,Дюймовочка", "['свет', 'солнце']"),
    ("Ласточка,Крот", "[]"),
    ("Мышь,Крот", "['деньги', 'уют']"),
    ("Крот,Мышь", "['деньги', 'уют']"),
    ("Дюймовочка,Принц эльфов", "['цветы']")
]

for inp, out in test_data:
    with Tester(inp.split(",")) as t:
        get_info()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача B5. Dull numbers

Напишите функцию `test()`, которая принимает на вход список строк, которые могут быть явно сконвертированы в целые неотрицательные числа, и возвращает список из логических значений (тип `boolean`) следующего вида:

* элемент равен `True`, если сумма цифр числа чётная;
* элемент равен `False` – в противном случае.


**Пример работы программы**

Запуск функции: 

    test(['23', '56', '123', '99'])

Результат:

    [False, False, True, True]
    
Запуск функции: 

    test(['101', '77', '88'])

Результат:

    [True, True, True]

```python
### YOUR CODE HERE ###
```

```python
assert test(['23', '56', '123', '99']) == [False, False, True, True], "Неверно"
assert test(['101', '77', '88']) == [True, True, True], "Неверно"
assert test(['101', '77', '88', '155', '12345', '8889']) == [True, True, True, False, False, False], "Неверно"
assert test(['1234', '00', '907654213', '1111']) == [True, True, False, True], "Неверно"
```

### Задача B6. Рейтинги мультфильмов

Необходимо посчитать средний рейтинг мультфильма в зависимости от его типа: рисованный или кукольный.

Напишите функцию `get_average()`, которая **принимает на вход** словарь и строку:

* Ключами словаря являются названия мультфильмов, а значениями — словари с их характеристиками (тип, рейтинг). Гарантируется, что в словаре есть информация минимум об одном мультфильме. Гарантируется, что все вложенные словари состоят из одинакового числа элементов. 
* Строка — тип мультфильма в словаре (рисованный, кукольный), соответствующем одному мультфильму, средний рейтинг по которому нужно посчитать.

И **возвращает** целое число — средний рейтинг, посчитанный по всем мультфильмам выбранного типа (сумму значений разделить на количество мультфильмов), *округленный до десятых*.

**Примеры работы программы**

Запуск функции:

    mults = {
        "Малыш и Карлсон" : {"type" : "рисованный", "rating" : 8.5}, 
        "Котенок с улицы Лизюкова" : {"type" : "рисованный", 
                           "rating" : 8.0},
        "Как стать большим" : {"type" : "кукольный", "rating" : 8.1},
        "Снежная королева" : {"type" : "рисованный", "rating" : 8},
        "Приключения капитана Врунгеля" : {"type" : "рисованный", 
                           "rating" : 8.2},
        "Буренка из Маслёнкино" : {"type" : "кукольный", "rating" : 6.9},
        "Сладкая сказка" : {"type" : "кукольный", "rating" : 7.3},
        "Жил-был пес" : {"type" : "рисованный", "rating" : 9.2},
    }

    get_average(mults, "рисованный")
    
    
Результат: 

    8
    
Запуск функции:

    mults = {
        "Малыш и Карлсон" : {"type" : "рисованный", "rating" : 8.5}, 
        "Котенок с улицы Лизюкова" : {"type" : "рисованный", 
                                "rating" : 8.0},
        "Как стать большим" : {"type" : "кукольный", "rating" : 8.1},
        "Снежная королева" : {"type" : "рисованный", "rating" : 8},
        "Буренка из Маслёнкино" : {"type" : "кукольный", "rating" : 6.9},
        "Жил-был пес" : {"type" : "рисованный", "rating" : 9.2},
    }

    get_average(mults, "кукольный")
    
    
Результат: 

    7

```python
### YOUR CODE HERE ###
```

```python
mults1 = {
    "Малыш и Карлсон" : {"type" : "рисованный", "rating" : 8.5}, 
    "Котенок с улицы Лизюкова" : {"type" : "рисованный", 
                       "rating" : 8.0},
    "Как стать большим" : {"type" : "кукольный", "rating" : 8.1},
    "Снежная королева" : {"type" : "рисованный", "rating" : 8},
    "Приключения капитана Врунгеля" : {"type" : "рисованный", 
                       "rating" : 8.2},
    "Буренка из Маслёнкино" : {"type" : "кукольный", "rating" : 6.9},
    "Сладкая сказка" : {"type" : "кукольный", "rating" : 7.3},
    "Жил-был пес" : {"type" : "рисованный", "rating" : 9.2},
}

mults2 = {
        "Малыш и Карлсон" : {"type" : "рисованный", "rating" : 8.5}, 
        "Котенок с улицы Лизюкова" : {"type" : "рисованный", "rating" : 8.0},
        "Как стать большим" : {"type" : "кукольный", "rating" : 8.1},
        "Снежная королева" : {"type" : "рисованный", "rating" : 8},
        "Буренка из Маслёнкино" : {"type" : "кукольный", "rating" : 6.9},
        "Жил-был пес" : {"type" : "рисованный", "rating" : 9.2},
    }

assert get_average(mults1, "кукольный") == 7.4, "Неверный ответ"
assert get_average(mults2, "кукольный") == 7.5, "Неверный ответ"
assert get_average(mults1, "рисованный") == 8.4, "Неверный ответ"
assert get_average(mults2, "рисованный") == 8.4, "Неверный ответ"
```

### Задача В7. Gradebook

Напишите программу, которая запрашивает с клавиатуры целое число $n$, затем $n$ раз запрашивает имя студента (одно слово) и перечень его оценок через пробел и выводит на экран список кортежей, где каждый кортеж состоит из имени студента и его среднего балла, округлённого до целых.

**Пример работы программы**

Ввод:

    3
    Ann 5 6 7 8
    Ben 10 10 0 0
    Fred 7 6 8 5

Вывод:

    [('Ann', 6), ('Ben', 5), ('Fred', 6)]
    
Ввод:

    4
    Мария 9 8 7
    Борис 10 8
    Екатерина 5 6 4 9 10
    Владимир 10 9 10

Вывод:

    [('Мария', 8), ('Борис', 9), ('Екатерина', 7), ('Владимир', 10)]

```python
def gradebook():
    
    ### YOUR CODE HERE ###
    
gradebook()
```

```python
test_data = [("4,Мария 9 8 7,Борис 10 8,Екатерина 5 6 4 9 10,Владимир 10 9 10",
             "[('Мария', 8), ('Борис', 9), ('Екатерина', 7), ('Владимир', 10)]"),
            ("1,Dan 2 5 8", "[('Dan', 5)]"),
            ("2,A 10 7 4,B 4 6 10 9", "[('A', 7), ('B', 7)]"),
            ("3,Ann 5 6 7 8,Ben 10 10 0 0,Fred 7 6 8 5", 
             "[('Ann', 6), ('Ben', 5), ('Fred', 6)]")]

for inp, out in test_data:
    with Tester(inp.split(",")) as t:
        gradebook()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
```

### Задача В8. Vipera evanesco!

Напишите программу, которая запрашивает с клавиатуры некоторый текст, заменяет в нём названия рептилий на `-` и выводит на экран полученный текст. Названия рептилий, которые допускаются при вводе, сохранены в списке `rept`:

    rept = ["snake", "python", "anaconda", "cobra", "taipan"]
    
**NB.** Ваша программа должна заменять названия рептилий, стоящих как в единственном, так и во множественном числе. Специальных библиотек для работы с частями речи в данном случае не требуется, достаточно учитывать слова из `rept` и те же слова с дописанным окончанием `s`.

Во время выполнения этой задачи нельзя использовать условные конструкции.

**Примеры работы программы**

Ввод:

    python is waiting for anaconda
    
Вывод:

    - is waiting for -
    
Ввод:

    just a racoon
    
Вывод:

    just a racoon
    
Ввод:

    we have no cobras here
    
Вывод:

    we have no - here

```python
def evanesco():
    
    rept = ["snake", "python", "anaconda", 
           "cobra", "taipan"]
    
    ### YOUR CODE HERE

evanesco()
```

```python
test_data = [("python", "-"), 
             ("no pythons", "no -"),
             ("python is waiting for anacondas", 
              "- is waiting for -"),
             ("just a racoon", "just a racoon"),
             ("we have no cobras here", "we have no - here"),
             ("snakes in the box", "- in the box"),
             ("is anaconda a snake?", "is - a -?"),
             ("taipans in India", "- in India"),
             ("taipan", "-")]

for inp, out in test_data:
    with Tester([inp]) as t:
        evanesco()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, была введена строка " + inp
```

### Задача C1. Back to school

Напишите функцию `eq_solve()`, которая принимает на вход строку с квадратным уравнением (без пробелов между знаками и со всеми коэффициентами, записанными явно, включая 1 и -1) и решает квадратное уравнение:

* если дискриминант отрицательный, возвращает пустой список;
* если дискриминант равен нулю, возвращает список из одного решения;
* если дискриминант больше нуля, возвращает список из двух решений, упорядоченных по возрастанию.

Считайте, что квадратное уравнение уже приведено к стандартному виду:

$$
ax^2+bx+c=0
$$

При решении нельзя использовать библиотечные функции (кроме функции для извлечения квадратного корня) и регулярные выражения.

**Примеры работы функции**

Применение:

    eq_solve('1x^2+6x+9=0')


Результат:

    [-3.0]

Применение:

    eq_solve('-2x^2-6x+9=0')

Результат:

    [-4.098076211353316, 1.098076211353316]

Применение:

    eq_solve('3x^2-1x+10=0')


Результат:

    []

```python
### YOUR CODE HERE
```

```python
from math import isclose

assert isclose(eq_solve('1x^2+6x+9=0')[0], -3.0), "Неверное решение при D = 0"
assert isclose(eq_solve('-2x^2-6x+9=0')[0], -4.098076211353316), "Неверное решение при D > 0"
assert isclose(eq_solve('-2x^2-6x+9=0')[1], 1.098076211353316), "Неверное решение при D > 0"
assert len(eq_solve('3x^2-1x+10=0')) == 0, "Неверное решение при D < 0"
assert isclose(eq_solve('17x^2-5x-7=0')[0], -0.5112655672235116), "Неверное решение при D > 0"
assert isclose(eq_solve('17x^2-5x-7=0')[1], 0.8053832142823351), "Неверное решение при D > 0"
assert isclose(eq_solve('-10x^2+25x-4=0')[0], 0.1718070673576088), "Неверное решение при D > 0"
assert isclose(eq_solve('-10x^2+25x-4=0')[1], 2.3281929326423914), "Неверное решение при D > 0"
assert len(eq_solve('9x^2+4x+4=0')) == 0, "Неверное решение при D < 0"
```

### Задача C2. Занимательная экономика

В словаре `animals` сохранены названия различных животных (включая птиц и насекомых) и все формы этих названий (все падежи в единственном и множественном числе).

Напишите функцию `animals_book()`, которая принимает на вход текст и возвращает список пар *название* – *число упоминаний* всех животных из `animals` в этом тексте, причем пары отсортированы по названию животного в алфавитном порядке. Упоминание животного – упоминание слова в любой форме, то есть любом падеже и числе.

Текст для отработки и проверки работы функции сохранен в переменной `song`, это шуточная песня о деньгах из фильма «Тим Талер, или Проданный смех» (Беларусьфильм, 1981). 

Образец результата применения функции `animals_book()` к тексту `song`:

    [('аист', 0), ('бегемот', 2), ('гусь', 1), ('дятел', 2), ('енот', 0), ('жираф', 0), 
     ('заяц', 2), ('кенгуру', 2), ('коза', 2), ('комар', 2), ('лягушка', 1), ('медведь', 1), 
     ('сова', 0)]

```python
animals = {
    "аист" : ["аист", "аисты", "аиста", "аистов", "аисту", "аистам", "аиста", "аистов", 
              "аистом", "аистами", "аисте", "аистах"],
    "гусь" : ["гусь", "гуси", "гуся", "гусей", "гусю", "гусям", "гуся", "гусей", "гусем", 
              "гусями", "гусе", "гусях"],
    "дятел" : ["дятел", "дятлы", "дятла", "дятлов", "дятлу", "дятлам", "дятла", "дятлов", 
               "дятлом", "дятлами", "дятле", "дятлах"],
    "сова" : ["сова", "совы", "совы", "сов", "сове", "совам", "сову", "сов", 
              "совой", "совами", "сове", "совах"],
    "бегемот" : ["бегемот", "бегемоты", "бегемота", "бегемотов", "бегемоту", "бегемотам", 
                 "бегемота", "бегемотов", "бегемотом", "бегемотами", "бегемоте", "бегемотах"],
    "енот" : ["енот", "еноты", "енота", "енотов", "еноту", "енотам", "енота", "енотов", 
              "енотом", "енотами", "еноте", "енотах"],
    "жираф" : ["жираф", "жирафы", "жирафа", "жирафов", "жирафу", "жирафам", 
               "жирафа", "жирафов", "жирафом", "жирафами", "жирафе", "жирафах"],
    "заяц" : ["заяц", "зайцы", "зайца", "зайцев", "зайцу", "зайцам", "зайца", "зайцев", 
              "зайцем", "зайцами", "зайце", "зайцах"],
    "коза" : ["коза", "козы", "козы", "коз", "козу", "коз", "козе", "козам", 
              "козой", "козами", "козе", "козах"],
    "кенгуру" : ["кенгуру", "кенгуру", "кенгуру", "кенгуру", "кенгуру", "кенгуру", "кенгуру",
                 "кенгуру", "кенгуру", "кенгуру", "кенгуру", "кенгуру"],
    "лягушка" : ["лягушка", "лягушки", "лягушки", "лягушек", "лягушке", "лягушкам", 
                 "лягушку", "лягушек", "лягушкой", "лягушками", "лягушке", "лягушках"], 
    "медведь" : ["медведь", "медведи", "медведя", "медведей", "медведю", "медведям", 
                 "медведя", "медведей", "медведем", "медведями", "медведе", "медведях"],
    "комар" : ["комар", "комары", "комара", "комаров", "комару", "комарам", 
               "комара", "комаров", "комаром", "комарами", "комаре", "комарах"]
}

song = """
Почему-то, отчего-то
Так устроен белый свет.
Денег нет у бегемота
И у зайца тоже нет.
И без денег солнце всходит
И прохладу дарит лес,
И лягушка на болоте
Комаров бесплатных ест,
Комаров бесплатных ест...

Для чего, ответьте,
Почему на свете
Лишь царям природы деньги нужны?
Ни коза, ни дятел
Никогда не платят,
Но платить все время люди должны.
Платить, платить за все,
Платить, платить должны.
Ни коза, ни дятел
Никогда не платят,
Но платить все время люди должны.

Кто когда придумал это,
Я ответить не берусь,
Но без авиабилета
В облаках летает гусь.
И река бесплатно берег
Умывает поутру,
И, конечно, не для денег
Есть карман у кенгуру,
Есть карман у кенгуру.

Почему-то, отчего-то
Так устроен белый свет.
Денег нет у бегемота
И у зайца тоже нет.
Ни копейки у медведя
Нет и не было вовек.
Человек придумал деньги,
Вот и платит человек,
Вот и платит человек.
"""
```

```python
### YOUR CODE HERE ###
```

*Отсылки в названиях и сюжетах задач*

*A2, B4. м/ф «Дюймовочка» (Союзмультфильм, 1964)*<br>
*A3. м/ф «Паровозик из Ромашкова» (Союзмультфильм, 1967)*<br>
*A6. песня "We Are the Champions" (Queen, 1977)*<br>
*B1. м/ф «Дом для Кузьки» (Экран, 1984)*<br>
*B2. м/ф «Чип и Дейл спешат на помощь» (Дисней, 1990)*<br>
*B3. м/ф «Тайна третьей планеты» (Союзмультфильм, 1981)*<br>
*B8. фильм «Гарри Поттер и Тайная комната» (Warner Bros. Pictures, 2002)*